# <span style="color:teal;">CIS 211 Project 3:  OOP</span>

##### Due 11:00 P.M. January 27, 2017

##### Reading:  M&R 10.1 -- 10.4

Our goal this week is to gain some experience writing class definitions for some simple objects.  The classes are part of a solar system simulation.  This week we'll define the objects used in the simulation.  Next week we'll implement the simulation itself, and later in the term we'll add visualization. 

###  <span style="color:teal">Part 1: Vectors (50 points)</span>

Define a class named Vector.  Instances of this class will be Euclidan vectors (https://en.wikipedia.org/wiki/Euclidean_vector).  

An instance of this class will have three attributes, representing $x$, $y$, and $z$ coordinates in space.  When a Vector object is created the constructor should be passed three numbers to use as the initial values of the coordinates.

Your class should also have:
* accessor functions named `x`, `y`, and `z`, which return the current value of the specified coordinate
* a `__repr__` function that displays a Vector as a tuple of 3 numbers; use `"%.3g"` in the format statement so coordinates are shown with 3 significant digits
* methods that implement `==`, `+`, `-`, and `*` operations (see below)
* a method named `norm` that computes the length, or magnitude, of the vector, defined by the equation $$ \lVert \, v \, \rVert = \sqrt{x^2 + y^2 + z^2} $$
* a method named `clear` that sets all coordinates to 0

To see if two vectors are equal, simply see if their $x$, $y$, and $z$ components are the same.  This will almost certainly fail (due to roundoff errors) when the coordinates are of the size used in the solar system simulation, but it will help test the other operations, and we'll use small numbers in these tests.

To add or subtract two vectors, make a new vector that has the sum or difference of the components of two existing vectors.

The method that implemements multiplication should implement scalar multiplication: the second operand should be an integer or a float, and the result is a new vector where all components are multiplied by the scalar.

Example:
<pre>
>>> v1 = Vector(3, 5, 0)
>>> v2 = Vector(1, 1, 4)

>>> v1.x()
3

>>> v1.y()
5

>>> v1.norm()
5.830951894845301

>>> v1 + v2
(4,6,4)

>>> v1 - v2
(2,4,-4)

>>> v2 * 3
(3,3,12)

>>> v1 + v2 == Vector(4,6,4)
True

>>> v1.clear()
>>> v1
(0,0,0)
</pre>

##### <span style="color:red">Code:</span>

<!-- grade_id: vector, points: 20, grade: true, solution: true -->
```python
from math import sqrt

class Vector:
    """
    A Vector is a 3-tuple of (x,y,z) coordinates.
    """
    def __init__(self, x, y, z):
        'Creates a new vector with coordinate x, y, and z'
        self._x = x
        self._y = y
        self._z = z
        
    def __repr__(self):
        return '(%.3g,%.3g,%.3g)'%(self._x, self._y, self._z)
    
    def x(self):
        'Return the x coordinate'
        return self._x
    
    def y(self):
        'Return the y coordinate'
        return self._y
        
    def z(self):
        'Return the z coordinate'
        return self._z
    
    def norm(self):
        'Return the magnitude of the vector'
        return sqrt(self._x**2 + self._y**2 + self._z**2)
    
    def clear(self):
        'Sets all coordinates to 0'
        self._x = 0
        self._y = 0
        self._z = 0
        return Vector(self._x, self._y, self._z)
        
    def __add__(self, other):
        return Vector(self._x + other._x, self._y + other._y, self._z + other._z)
    
    def __sub__(self, other):
        return Vector(self._x - other._x, self._y - other._y, self._z - other._z)
    
    def __mul__(self, n):
        return Vector(self._x * n, self._y * n, self._z * n)
    
    def __eq__(self, other):
        return (self._x == other._x) and (self._y == other._y) and (self._z == other._z)

    
    
```

##### <span style="color:red">Tests:</span>

**Optional** &nbsp; If you want to do your own tests use the code cell below to create and test objects.  You can add additional cells here if you want.

##### <span style="color:red">Autograder Tests:</span>

<!-- grade_id: vec_setup, points: 2, grade: true, solution: false -->
```python
v1 = Vector(1, 1, 1)
v2 = Vector(2, 2, 2)
v3 = Vector(3, 3, 3)
```

<!-- grade_id: vec_repr, points: 2, grade: true, solution: false -->
```python
assert str(v1) == '(1,1,1)'
assert str(v2) == '(2,2,2)'
```

<!-- grade_id: vec_eq, points: 2, grade: true, solution: false -->
```python
assert Vector(1, 2, 3) == Vector(1, 2, 3)
assert Vector(1, 2, 3) != Vector(1.001, 2.001, 3.001)
```

<!-- grade_id: sum_test, points: 2, grade: true, solution: false -->
```python
assert v1 + v2 == v3
```

<!-- grade_id: diff_test, points: 3, grade: true, solution: false -->
```python
assert v3 - v2 == v1
```

<!-- grade_id: scale_test, points: 3, grade: true, solution: false -->
```python
assert v3 == v1 * 3
```

<!-- grade_id: mag_test, points: 3, grade: true, solution: false -->
```python
assert round(v2.norm(), 10) == round(sqrt(12), 10)
```

<!-- grade_id: clear_test, points: 3, grade: true, solution: false -->
```python
v1.clear()
assert v1 == Vector(0,0,0)
```

##### <span style="color:red">Documentation:</span>

<!-- grade_id: doc_1, points: 10, grade: true, solution: true -->
Class Vector has a constructor that initializes x, y, and z. Within the class, there are methods that return the value of all objects, a 'norm' function that returns the magnitude of the vector, a function that clears the Vector itself, and functions that return the sum, difference, and product of the Vectors, and a final function that gives a boolean return if the x, y, and z are equal to each other.

###  <span style="color:teal">Part 2: Celestial Bodies (50 points)</span>

Define a class named `Body` that will be used to represent the Sun and planets in the solar system simulation.

Each instance of this class should have the following attributes:
* mass, represented by a single floating point number
* position, velocity, and force, each of which is a 3D vector
* a name

Define the constructor so that all attributes are optional.  The default name is None, mass is 0, and position and velocity are (0,0,0).  The force vector should always be initialized to (0,0,0).

The representation string for a Body should include the name, if it's not None, and then the mass, position, and velocity:

<pre>
>>> Body()
0kg (0,0,0) (0,0,0)

>>> ep = Vector(-2.700743E+10, 1.446007E+11, 9686451)
>>> Body(name = 'earth', mass = 5.9736E+24, position = ep)
earth: 5.97e+24kg (-2.7e+10,1.45e+11,9.69e+06) (0,0,0)
</pre>

Your class should also include the following:
* accessor functions called `name`, `mass`, `position`, `velocity`, and `force` that return the corresponding attribute
* a method called `direction` which takes another Body object as a parameter and returns a vector that "points at" the other Body (see example below)
* a method named `clear_force` that sets the force vector to (0,0,0) by calling that vector's `clear` method
* an `add_force` method that takes another Body as a parameter and updates the force vector using the equation shown below
* a method called `move` that will use the current value of the force vector to update the object's position using the algorithm outlined below.

This example illustrates the `direction` method.  Notice how the direction from `b1` to `b2` is the same size but points the other way from the direction from `b2` to `b1`.
<pre>
>>> b1 = Body(position = Vector(0,1,0))
>>> b2 = Body(position = Vector(1,0,0))

>>> b1.direction(b2)
(1,-1,0)

>>> b2.direction(b1)
(-1,1,0)
</pre>

To implement `add_force`, first compute the force pulling a body A toward another body B. Let $\vec{d}$ be the direction from A to B. Then a vector $\vec{f}$ that defines the force is

$$
\vec{f} = \frac{\vec{d} \times m_\mathrm{B}}{{\lVert \, \vec{d} \, \rVert}^3}
$$

where $m_\mathrm{B}$ is the mass of body B.  Add $\vec{f}$ to the force vector in the object that represents A.

The `move` method will take a parameter named `dt` which is a time step size.
To move a body A use the accumulated forces created by previous calls to `add_force`.
First compute a vector that represents the acceleration of A:

$$
\vec{a} = G \times \vec{f}_\mathrm{A} 
$$

where $G$ is the universal gravitational constant and $\vec{f}_\mathrm{A}$ is A's force vector.

Then update A's velocity vector:
$
\vec{v}_\mathrm{A} = \vec{v}_\mathrm{A} + \vec{a} \times \mathtt{dt}
$

Finally, update A's position vector:
$
\vec{p}_\mathrm{A} =  \vec{p}_\mathrm{A} + \vec{v}_\mathrm{A} \times \mathtt{dt}
$



##### <span style="color:red">Code:</span>

Put your class definition in the following code cell.  Note the gravitational constant $G$ has been defined for you.

<!-- grade_id: body, points: 20, grade: true, solution: true -->
```python
G = 6.67E-11

class Body:
    """
    A Body object represents the state of a celestial body.  A body has mass 
    (a scalar), position (a vector), and velocity (a vector).  A third vector, 
    named force, is used when calculating forces acting on a body.  An
    optional name can be attached to use in debugging.
    """
    
    def __init__(self, mass = 0, position = Vector(0,0,0), velocity = Vector(0,0,0), name = None):
        """
        Create a new Body object with the specified mass (a scalar), position (a vector), 
        and velocity (another vector).  A fourth argument is an optional name for the body.
       """
        self._name = name
        self._mass = mass
        self._pos = position
        self._vel = velocity
        self._force = Vector(0,0,0)

    def __repr__(self):
        if self._name == None:
            return '{:.3g}kg {} {}'.format(self._mass, self._pos, self._vel)
        return'{}: {:.3g}kg {} {}'.format(self._name, self._mass, self._pos, self._vel)
    
    def name(self):
        return self._name

    def mass(self):
        return self._mass

    def position(self):
        return self._pos

    def velocity(self):
        return self._vel

    def force(self):
        return self._force
    
    def direction(self, other):
        return other.position() - self.position()
    
    def add_force(self, other):
        dcubed = self.direction(other).norm() * self.direction(other).norm() * self.direction(other).norm()
        self._force = self._force + (self.direction(other) * (other.mass()/dcubed))
        return self._force
        
    def clear_force(self):
        self._force = Vector(0,0,0)
        return self._force
         
    def move(self, dt):
        acc = self.force() * G
        self._vel = self._vel + (acc * dt)
        self._pos = self._pos + (self._vel * dt)
```

##### <span style="color:red">Tests:</span>

**Optional** &nbsp; If you want to do your own tests use the code cell below to create and test objects.  You can add additional cells here if you want.

```python

```

##### <span style="color:red">Autograder Tests:</span>

<!-- grade_id: body_setup, points: 3, grade: true, solution: false -->
```python
melon = Body(name='melon', mass=3.0, position=Vector(0,6371000,0))
earth = Body(name='earth', mass=5.9736E+24)
```

<!-- grade_id: body_repr, points: 3, grade: true, solution: false -->
```python
assert str(melon) == "melon: 3kg (0,6.37e+06,0) (0,0,0)"
assert str(earth) == "earth: 5.97e+24kg (0,0,0) (0,0,0)"
```

<!-- grade_id: dir_test, points: 3, grade: true, solution: false -->
```python
assert earth.direction(melon) == Vector(0,6371000,0)
assert melon.direction(earth) == Vector(0,-6371000,0)
```

<!-- grade_id: add_force_test, points: 4, grade: true, solution: false -->
```python
melon.add_force(earth)
assert -1.48e+11 < melon.force().y() < -1.47e+11
```

<!-- grade_id: move_test, points: 4, grade: true, solution: false -->
```python
melon.move(1)
assert 6370990 < melon.position().y() < 6371000
```

<!-- grade_id: clear_force_test, points: 3, grade: true, solution: false -->
```python
melon.clear_force()
assert melon.force() == Vector(0,0,0)
```

##### <span style="color:red">Documentation:</span>

<!-- grade_id: doc_2, points: 10, grade: true, solution: true -->
Class Body that constructs the name (unless is None), mass, position, velocity, and force of a celestial body. Contains methods of the direction of the body, the force added to an object, and to move the body. There is also a clear_force method that returns a clear vector.