### <span style="color:teal;">CIS 211 Project 5:  Building GUIs with Tk</span>

##### Due 11:00 P.M. Friday Feb 10

Reading:
* [TkDocs - Tk Tutorial](http://www.tkdocs.com/tutorial/index.html)
* [Tkinter Tutorial](http://www.python-course.eu/python_tkinter.php) from `python-course.eu`

### <span style="color:teal">Setup</span>

Execute the following code cell to tell IPython you want to create windows with Tk and to import the `tkinter` library.

```python
%gui tk
import tkinter as tk
```

<span style="color:red">**Important:**</span> &nbsp; The autograder tests assume `tkinter` has been imported as shown above.  You should follow the same convention, and refer to widget classes as `tk.Label`, `tk.Entry`, and so on.

<span style="color:red">**Important:**</span> &nbsp; The code cells where you define your GUIs contain assignment statements that create global variables.  For example, the first GUI is assigned to a variable named `hello_app`.  **Do not change these assignment statements**.  The autograder test cells will use these variables to access parts of your GUIs.

### <span style="color:teal">1. &nbsp; Hello, World (10 points)</span>

The code cell for this project has an assignment statement that creates a top level Tk application named `hello_app`.  

Add statements to the code cell so the program has a label that says "Hello, World":

<img src="http://www.cs.uoregon.edu/Classes/15S/cis211/images/hello_gui.png"/>

The example above uses 24-point Helvetica.  You can use the default text, but feel free to experiment with other fonts, font sizes, colors, *etc*.

The goal for this part of the project is just to make sure you have all the necessary software installed and can create a simple application.  There is no documentation requirement for this part.

##### <span style="color:red">Code:</span>

<!-- grade_id: hello, points: 4, grade: true, solution: true -->
```python
hello_app = tk.Tk()
def say_hello():
    w = tk.Label(hello_app, text="Hello, World!")
    w.configure(font = ('Helvetica',24))
    w.pack()
```

```python
say_hello()
```

##### <span style="color:red">Autograder Tests:</span>

<!-- grade_id: hello_test_1, points: 2, grade: true, solution: false -->
```python
# The GUI should have a single widget
widgets = list(hello_app.children.values())
assert len(widgets) == 1
```

<!-- grade_id: hello_test_2, points: 2, grade: true, solution: false -->
```python
# Check to see if the widget is a Label
widgets = list(hello_app.children.values())
assert isinstance(widgets[0], tk.Label)
```

<!-- grade_id: hello_test_3, points: 2, grade: true, solution: false -->
```python
# The text in the label should start with "hello" (upper or lower case)
widgets = list(hello_app.children.values())
assert widgets[0]['text'].lower().startswith('hello')
```

### <span style="color:teal">2. &nbsp; Payment Calculator (50 points)</span>

Create a program that computes the monthly payment on a loan. The payment is a function of the loan amount, the annual interest rate, and the number of years to pay off the loan. To calculate the payment, first compute two values named $r$ and $p$, defined by these equations:

$ r = ( \textrm{rate} \; / \; 100 ) \; / \; 12 $

$ p = 12 \times \textrm{years} $

The monthly payment is then

$$ \textrm{payment} = \frac{ r \times \textrm{amount} }{ 1 - (1 + r)^{-p} } $$

For example, if you take out a 30-year loan for \$150,000 at an annual rate of 4.5% your payment would be \$760.03 per month.

The GUI for this application should have four text entry boxes, and each box should have a label to its left that explains what the box holds.  The first three boxes are for the user to enter the loan amount, the interest rate, and the number of years to repay the loan. 

The GUI should have a button labeled “compute payment”. When the user clicks this button, your program should read the values of the three parameters, use the equations shown above to calculate the monthly payment, and display the payment amount in the fourth text entry box.

<img src="http://www.cs.uoregon.edu/Classes/15S/cis211/images/payment_gui.png"/>

The code cell for this program has the outline for the program.  First fill in the body of the function named `monthly_payment` and make sure it computes the correct payment value given an amount, rate, and loan period.  When your write the button callback function have it call the `monthly_payment` function.

Next add statements that create the labels, boxes, and button for the application.  When you create the Entry and Button widgets, make sure you save them in global variables with the following names so the autograder tests will succeed:
* `amount`
* `rate`
* `years`
* `payment`
* `button`

##### <span style="color:red">Code:</span>

<!-- grade_id: payment, points: 15, grade: true, solution: true -->
```python
def monthly_payment(amt, pct, yrs):
    r = (float(pct)/100)/12
    p = 12*float(yrs)
    pay = (r*float(amt))/(1 - (1 + r)**(-p))
    return pay
    
def compute_payment_cb():
    amt = amount.get()
    pct = rate.get()
    yrs = years.get()
    monthpay = monthly_payment(amt, pct, yrs)
    payment.insert(0,'${:.2f}'.format(monthpay))
    
    
    
payment_app = tk.Tk()

##loan payment
loanlabel = tk.Label(payment_app, text = "Loan Amount:")
loanlabel.grid(row=0, column=0, sticky=tk.W, padx = 20, pady = 30)

amount = tk.Entry(payment_app)
amount.grid(row=0, column=1, padx = 20)

##interest rate
intlabel = tk.Label(payment_app, text = "Interest Rate:")
intlabel.grid(row=1, column=0, sticky=tk.W, padx = 20, pady = 30)

rate = tk.Entry(payment_app)
rate.grid(row=1, column=1, padx = 20)

##loan period
lplabel = tk.Label(payment_app, text = "Loan Period:")
lplabel.grid(row=2, column=0, sticky=tk.W, padx = 20, pady = 30)

years = tk.Entry(payment_app)
years.grid(row=2, column=1, padx = 20)

##payment
paylabel = tk.Label(payment_app, text = "Payment:")
paylabel.grid(row=3, column=0, sticky=tk.W, padx = 20, pady = 30)

payment = tk.Entry(payment_app)
payment.grid(row=3, column=1, padx = 20)

#button
button = tk.Button(payment_app, text='Compute Payment', command=compute_payment_cb)
button.grid(row=4, column=0, columnspan=2, pady = 30)



```

##### <span style="color:red">Autograder Tests:</span>

<!-- grade_id: payment_test_1, points: 5, grade: true, solution: false -->
```python
# Test the callback function to make sure it is computing payments accurately
assert round(monthly_payment(20000, 3.0, 6), 2) == 303.87
```

<!-- grade_id: payment_test_2, points: 10, grade: true, solution: false -->
```python
# Verify the number of widgets of each type
counts = { tk.Label: 0, tk.Entry: 0, tk.Button: 0 }

for x in payment_app.children.values():
    counts[type(x)] += 1
    
assert counts == { tk.Label: 4, tk.Entry: 4, tk.Button: 1 }
```

<!-- grade_id: payment_test_3, points: 10, grade: true, solution: false -->
```python
# Store values in the GUI, click the button, check the result
amount.insert(0, '20000')
rate.insert(0, '3.0')
years.insert(0, '6')
button.invoke()

assert payment.get() == '$303.87'
```

##### <span style="color:red">Documentation:</span>

<!-- grade_id: doc_2, points: 10, grade: true, solution: true -->
The code for part 2 contains a monthly_payment function that takes in 3 arguments amt, pct, and yrs for the amount interest rate and years for the calculation. In the compute_payment_cb it initializes the variables listed earlier using the global variables from the input taken from the box widget. I then use an object called monthlypay to call back the monthly payment calculation and insert it into the payment box widget as a float to 2 decimal places.

### <span style="color:teal">3. &nbsp; Canvas (40 points)</span>

Fill in the application named `canvas_app`.  The application will have just one widget, a Tk canvas with a size of $400 \times 400$ pixels. 

You also need to fill in the body of the function named `make_circle` that will draw a circle on your canvas.  The parameters are a reference to the canvas widget, three integers for the size and location of the circle, and a string with a color name.

Use the function to draw three circles on your canvas.  You can choose whatever size, location, and color you would like.

<img src="http://www.cs.uoregon.edu/Classes/16W/cis211/images/gui/three_circles.png"/>


##### <span style="color:red">Code:</span>

<!-- grade_id: canvas, points: 15, grade: true, solution: true -->
```python
def make_circle(x, y, radius, color):
    drawing.create_oval((x-radius,y-radius),(x+radius,y+radius),fill=color)
        

canvas_app = tk.Tk()
drawing = tk.Canvas(canvas_app, height=400, width=400)
drawing.pack()

```

```python
make_circle(60,80,50,"blue")
make_circle(100,200,10,"red")
make_circle(200,300,15,"yellow")
```

##### <span style="color:red">Autograder Tests:</span>

<!-- grade_id: canvas_test_1, points: 5, grade: true, solution: false -->
```python
# See if the GUI has a Canvas widget
widgets = list(canvas_app.children.values())
assert len(widgets) == 1 and isinstance(widgets[0], tk.Canvas)
```

<!-- grade_id: canvas_test_2, points: 10, grade: true, solution: false -->
```python
# See if the canvas has circles on it
widgets = list(canvas_app.children.values())
for n in widgets[0].find_all():
    assert widgets[0].type(n) == 'oval'
```

##### <span style="color:red">Documentation:</span>

<!-- grade_id: doc_3, points: 10, grade: true, solution: true -->
All I did for this code was to take in the first x,y coordinate and subtract it by the radius in order to get the top left corner of the circle, and then take that same x,y coordinate and add the radius to get the bottom right coordinate in order to make the full circle.